After you've completed the lab, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "" # put your full name here
COLLABORATORS = "" # put the name(s) of anyone you worked with here
```

---

# Lab 1: Getting started

**This assignment was distributed on Monday September 2nd. There is no due date for this assignment, but you should try to complete it and ask any questions in office hours or on Piazza before lab 2, which will be dsitributed on Monday September 9th.**

This is the first lab of ER131 and is intended to ensure that you have your Jupyter Hub and/or local Notebook environments configured correctly.

# Imports

These packages will be used over and over in the course.  In your own coding you don't *have* to import `numpy` as `np` (and so on), but it is what most people use, and using those abbreviations will improve how well your code ports over to others' environments.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import re
%matplotlib inline
```

# Exercises

## Part A

Write a function that returns a list of numbers, such that $ x_i=i^2 $,for $1\leq i \leq n$. 

* Don't worry about the case where $n < 0$.
* The assert statements below will be very picky on where spaces show up, and this narrows the range of solutions available to print out the list

Additionally, write the function so that it raises error `NotImplementedError` *unless* you give the function something to return *before* the `NotImplementedError` statement appears. You can use the following syntax to raise an error:

```python 
raise NotImplementedError()
```
If you finish early, try to come up with more than one way to code this.

<!-- grade_id: cell-9388c587d0e639ff, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """
    This function takes an integer n>0 as input.
    It computes the squares of numbers from 1 to n and returns a list,
    such that the ith element of the returned list equals i^2.
    """
    
    # YOUR CODE HERE
    
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$.  
Check that it does:

```python
squares(10)
```

Check that squares returns the correct output for several inputs.  

Note, the `assert` syntax below will throw `AssertionError`s if the boolean statements evaluate to `false` and do nothing if the statements are `true`. (You want your statements to be true.)

<!-- grade_id: squares-tests, points: 1, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

## Part B
Write a function that "cleans" a string by:
1.  making all letters lowercase,
2.  replacing any non-alphanumeric characters (not a-z or 0-9) with a space (hint: you'll have to use [regular expressions](https://docs.python.org/2/library/re.html) for this part), and
3.  trimming leading and ending spaces.

The function should take in a "raw" string as input and return a "clean" string.

<!-- grade_id: cell-b1fa3ce3a6566fbc, points: 0, grade: false, solution: true -->
```python
def clean(string):
    """Return cleaned string"""
    # YOUR CODE HERE

    raise NotImplementedError()
```

Your function should clean the address "#12 Oak Street " to "12 oak street":

```python
clean("#12 Oak Street")
```

Check that the function works as expected for several inputs.

<!-- grade_id: test-is-palindrome, points: 1, grade: true, solution: false -->
```python
assert clean("ER 131-lab section 102") == "er 131 lab section 102"
assert clean(" 08/28/2019 ") == "08 28 2019"
assert clean("$1500 DOLLARS ") == "1500 dollars"
```

## Part C

Run the following code to see a pretty plot.

```python
from mpl_toolkits.mplot3d import axes3d
import matplotlib.pyplot as plt


# Create a plot
u = np.linspace(1.5*np.pi, -1.5*np.pi, 100)
[x,y] = np.meshgrid(u, u)
z = np.cos( np.sqrt(x.flatten()**2 + y.flatten()**2) ).reshape(x.shape)


fig = plt.figure()
ax = fig.add_subplot(111, projection='3d')
# Plot a basic wireframe.
ax.plot_wireframe(x, y, z, rstride=10, cstride=10)
ax.view_init(elev=50., azim=30)
plt.savefig("figure1.pdf")
```